# SLU1 - Pandas 101: Exercise notebook

```python
import pandas as pd
```

In this notebook the following is tested:

- Pandas Series
- Pandas DataFrames
- Adding columns to a dataframe
- Printing the columns
- Load a dataset
- Preview a dataframe
- Make use of info, describe and shape

## Useful information

|    Country    | Population |   Capital  | Area (km^2) | Population Density (/km^2) | Main Religion | GDP (B) | Has Sovereign State? |
|:-------------:|:----------:|:----------:|:-----------:|:--------------------------:|:-------------:|:-------:|:--------------------:|
|    Denmark    |   5724456  | Copenhagen |    43094    |            129.5           |    Lutheran   |  306.7  |          No          |
|    Finland    |   5498211  |  Helsinki  |    338145   |            16.2            |    Lutheran   |  236.8  |          No          |
|    Iceland    |   335878   |  Reykjavík |    103000   |             3.2            |    Lutheran   |   20.0  |          No          |
|     Norway    |   5265158  |    Oslo    |    323802   |            16.1            |    Lutheran   |  370.4  |          No          |
|   Greenland   |    56483   |    Nuuk    |   2166086   |            0.028           |    Lutheran   |   2.22  |          Yes         |
| Faroe Islands |    49188   |  Tórshavn  |     1393    |            35.5            |    Lutheran   |   2.77  |          Yes         |

---

## Exercise 1: Series

In this first exercise the goal is to get used to creating series.

#### 1.1) Create a series for the countries

<!-- grade_id: cell-6e62f4a661edbaf0, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-e1f0144e6b6e35ff, points: 3, grade: true, solution: false -->
```python
assert isinstance(countries, pd.Series), "Should be of type pd.Series"
assert (countries == ["Denmark", "Finland", "Iceland", "Norway", "Greenland", "Faroe Islands"]).all()
```

#### 1.2) Create a series for the capitals (using countries as index)

<!-- grade_id: cell-7c2e90bd431de044, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-2a4b1b45d511997f, points: 2, grade: true, solution: false -->
```python
assert isinstance(capitals, pd.Series), "Should be of type pd.Series"
assert (capitals == ["Copenhagen", "Helsinki", "Reykjavík", "Oslo", "Nuuk", "Tórshavn"]).all()
assert (capitals.index == ["Denmark", "Finland", "Iceland", "Norway", "Greenland", "Faroe Islands"]).all()
```

---

## Exercise 2: DataFrames

In this exercise the goal is to create a simple DataFrame comprising the information provided in the "Useful Information" section.

#### 2.1) Create a DataFrame with countries/capital/population columns

<!-- grade_id: cell-1203e8c99879dd1a, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-59ff17beb4240534, points: 3, grade: true, solution: false -->
```python
assert isinstance(nordic_countries, pd.DataFrame), "Should be of type pd.DataFrame"
assert (nordic_countries["capital"] == ["Copenhagen", "Helsinki", "Reykjavík", "Oslo", "Nuuk", "Tórshavn"]).all()
assert (nordic_countries["countries"] == ["Denmark", "Finland", "Iceland", "Norway", "Greenland", "Faroe Islands"]).all()
assert (nordic_countries["population"] == [5724456, 5498211, 335878, 5265158, 56483, 49188]).all()
```

The expected output is a dataframe with columns named "capital", "countries" and "population", indexed from 0 to 5.

#### 2.2) Create a DataFrame with the capital/population columns, but with countries as row indexes

<!-- grade_id: cell-689992e365127432, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-aa4f3ae216111a48, points: 2, grade: true, solution: false -->
```python
assert isinstance(nordic_countries_ind, pd.DataFrame), "Should be of type pd.DataFrame"
assert (nordic_countries_ind["capital"] == ["Copenhagen", "Helsinki", "Reykjavík", "Oslo", "Nuuk", "Tórshavn"]).all()
assert (nordic_countries_ind["population"] == [5724456, 5498211, 335878, 5265158, 56483, 49188]).all()
assert (nordic_countries_ind.index == ["Denmark", "Finland", "Iceland", "Norway", "Greenland", "Faroe Islands"]).all()
```

The expected output is a dataframe with columns named "capital", "population", indexed with the correct countries.

#### 2.3) Going back to the 2.1 DataFrame, add columns with information concerning the gdp and the main religion (Do not create a dataframe from scratch)

<!-- grade_id: cell-ca9e66b218fd9ae1, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-81a422801f205db6, points: 2, grade: true, solution: false -->
```python
assert isinstance(nordic_countries, pd.DataFrame), "Should be of type pd.DataFrame"
assert (nordic_countries["gdp"] == [306.7, 236.8, 20.0, 370.4, 2.22, 2.77]).all()
assert (nordic_countries["mainReligion"] == ["Lutheran", "Lutheran", "Lutheran", "Lutheran", "Lutheran", "Lutheran"]).all()
assert set(nordic_countries.columns) == {"capital", "countries", "population", "gdp", "mainReligion"}
```

The expected output is a dataframe with columns "capital", "countries", "population", "gdp" and "mainReligion", indexed from 0 to 5.

## Exercise 3: Load the Iris csv data set

In this exercise you will load the iris data set ([source](https://gist.github.com/curran/a08a1080b88344b0c8a7)), one of the standard data sets for students of this field. You will then preview it and retrieve some information.

#### 3.1) Load the Iris data set

<!-- grade_id: cell-2a993c3e52e74d05, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-3d25559c7d603e24, points: 2, grade: true, solution: false -->
```python
assert isinstance(iris_df, pd.DataFrame), "Should be of type pd.DataFrame"
```

#### 3.2) Get general information about the dataframe

<!-- grade_id: cell-4e4fc6fed6166103, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

#### 3.3) Preview the top 10 entries

<!-- grade_id: cell-061a4e39b26e985f, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

#### 3.4) Preview the bottom 10 entries

<!-- grade_id: cell-98235f4828bd77e5, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

#### 3.5) Print the number of rows and the number of columns

<!-- grade_id: cell-0c5b18b72aed8fc2, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-c985768142805f1a, points: 2, grade: true, solution: false -->
```python
assert nr_rows == 150
assert nr_cols == 5
```

#### 3.6) Create a list with the variable (column) names

<!-- grade_id: cell-5cab778da54f035c, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-c4048ebebd4e4628, points: 2, grade: true, solution: false -->
```python
assert "sepal_width" in list_of_variables
assert "species" in list_of_variables
assert "sepal_length" in list_of_variables
```

The expected output is a list with "sepal_length", "sepal_width", "petal_length", "petal_width" and "species".

#### 3.7) Making use of describe() what is the mean value of the petal length?

<!-- grade_id: cell-48a24d2e60beee08, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-b619728041762ed8, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-91a22885dd4121c5, points: 2, grade: true, solution: false -->
```python
assert petal_length_mean < 3.8 and petal_length_mean > 3.7
```