# SLU1 - Pandas 101: Exercise notebook

First things first, import:


```python
import pandas as pd
import numpy as np
from hashlib import sha256
import json
```

In this notebook the following is tested:

- Pandas Series
- Pandas DataFrames
- Printing the columns
- Load a dataset
- Preview a dataframe
- Make use of info and describe

## Exercise 1: Series

In this first exercise the goal is to get used to creating series.

#### 1.1) Create a series with the name of the football club champions of 2019:

<!-- grade_id: cell-f0653e44790ad54d, points: 0, grade: false, solution: true -->
```python
# Create a series with the names of the football champions, using the order provided.
# Let pandas generate the index on its own. 
# Don't forget to delete the raise error

football_clubs_list = ['Barcelona', 'Benfica', 'Juventus', 'PSG', 'Manchester City']

# Create a series for the football_clubs, and call it football_clubs_series. 
# football_clubs_series = ...

# YOUR CODE HERE
raise NotImplementedError()
```

```python
# run this cell to see the results 
football_clubs_series
```

<!-- grade_id: cell-cf5318d87c48bfe0, points: 2, grade: true, solution: false -->
```python
assert len(football_clubs_series.index ) == 5 
assert ((football_clubs_series.values[3] == 'PSG') & (football_clubs_series.values[-1] =='Manchester City'))
assert (isinstance(football_clubs_series, pd.Series))
```

#### 1.2) Create a series for the countries of each football club (using the football clubs as index):

<!-- grade_id: cell-4cfa16d66278520b, points: 0, grade: false, solution: true -->
```python
football_clubs_list = ['Barcelona', 'Benfica', 'Juventus', 'PSG', 'Manchester City']

country_list = ['Spain', 'Portugal', 'Italy', 'France', 'England']

# Create a series which has the football club as index, and the respective country as values. 
# country_series = ...

# YOUR CODE HERE
raise NotImplementedError()
```

```python
# run this cell to see the results 
country_series
```

<!-- grade_id: cell-16e68b6e73803f03, points: 2, grade: true, solution: false -->
```python
assert (isinstance(country_series, pd.Series))
assert(len(country_series==5))
assert(country_series[-3] == 'Italy')
assert(country_series.index[3] == 'PSG')
```

#### 1.3) Get an numpy array with the the index of the given series.

<!-- grade_id: cell-e31640eaa1b1fdd5, points: 0, grade: false, solution: true -->
```python
series = pd.Series(['Spain', 'Portugal', 'Italy', 'France', 'England'], 
                          index=['Barcelona', 'Benfica', 'Juventus', 'PSG', 'Manchester City'])

# Get an array with the indexes of the serie.
# index_series = ...

# YOUR CODE HERE
raise NotImplementedError()
```

```python
# run this cell to see the results 
index_series
```

<!-- grade_id: cell-d1ad32e8fa3b6b7f, points: 1, grade: true, solution: false -->
```python
assert (isinstance(index_series, np.ndarray))
assert(len(index_series==5))
assert(index_series[2] == 'Juventus')
```

## Exercise 2: DataFrames

In this exercise the goal is to create a simple DataFrame.

#### 2.1) Create a DataFrame with FootballClub/Country/TotalMarketValue columns

<!-- grade_id: cell-7244c4ebd62cbd53, points: 0, grade: false, solution: true -->
```python
# Create a dataframe with 3 columns: FootballClub, Country, TotalMarketValue. 
# Let pandas create the index itself
# Make sure the columns are in the right order 
# winners_2019 must be the name of the final dataframe.

football_clubs_list = ['Barcelona', 'Benfica', 'Juventus', 'PSG', 'Manchester City']

country_list = ['Spain', 'Portugal', 'Italy', 'France', 'England']

totalmarketvalue_list = [1070.0, 365.50, 745.50, 922.25, 1170.0]

# in one line, build the dataframe, using a dictionary to tell it what should be the column names
# remember to give the columns the right names, or it won't pass the asserts
# winners_2019 = ...


# YOUR CODE HERE
raise NotImplementedError()
```

```python
# The expected output is a dataframe with columns named
# "FootballClub", "Country" and "TotalMarketValue", indexed from 0 to 4.
winners_2019
```

<!-- grade_id: cell-d51c1ab3368df679, points: 2, grade: true, solution: false -->
```python
assert(isinstance(winners_2019, pd.DataFrame))
assert(winners_2019.shape == (5, 3))
assert(winners_2019.iloc[3, 1] == 'France')
assert(winners_2019.columns[2] == 'TotalMarketValue')
```

#### 2.2) Create a DataFrame with the Country/TotalMarketValue columns, but with FootballClub as row indexes

<!-- grade_id: cell-7235c1b2a064e8a6, points: 0, grade: false, solution: true -->
```python
# Create a dataframe with columns Country and TotalMarketValue (in that order), but use FootballClub as index
# winners_2019_ind must be the name of the final dataframe.

football_clubs_list = ['Barcelona', 'Benfica', 'Juventus', 'PSG', 'Manchester City']

country_list = ['Spain', 'Portugal', 'Italy', 'France', 'England']

totalmarketvalue_list = [1070.0, 365.50, 745.50, 922.25, 1170.0]

# Use the lists to create Series for Country and TotalMarketValue
# Hint: use the argument index

# country_series = ...

# totalmarketvalue_series = ...

# Finally create the DataFrame, using a dictionary. This should only be 1 line of code. 
# winners_2019_ind = ...

# YOUR CODE HERE
raise NotImplementedError()
```

```python
# run this cell to see the results 
# The expected output is a dataframe with columns named "Country", "TotalMarketValue", indexed with the correct football clubs.
winners_2019_ind
```

<!-- grade_id: cell-9aac3d117a3c5976, points: 2, grade: true, solution: false -->
```python
assert(isinstance(winners_2019_ind, pd.DataFrame))
assert(winners_2019_ind.shape == (5, 2))
assert(winners_2019_ind.index[3] == 'PSG')
assert(winners_2019_ind.columns[1] == 'TotalMarketValue')
backup_winners_2019_ind = winners_2019_ind.copy()
```

## Exercise 3: Load the zomato dataset


In this exercise you will load the zomato data set.

You will then preview it and retrieve some information.

#### 3.1) Load the data set

<!-- grade_id: cell-18f79195c471e22b, points: 0, grade: false, solution: true -->
```python
# Load the zomato data set
# it is located at data/zomato_restaurants.csv
# zomato_df = ...

# YOUR CODE HERE
raise NotImplementedError()
```

```python
# print head it to get an idea of what you've just loaded 
zomato_df.head()
```

<!-- grade_id: cell-170ba41328e42eb0, points: 1, grade: true, solution: false -->
```python
assert(isinstance(zomato_df, pd.DataFrame))
assert(zomato_df.shape == (600, 14))
```

#### 3.2) Load the data set with an index defined

<!-- grade_id: cell-56561f9feb6b517e, points: 0, grade: false, solution: true -->
```python
# Load the zomato data set ans set the index the "id" column 
# (check the documentatio: https://pandas.pydata.org/pandas-docs/stable/reference/api/pandas.read_csv.html)
# it is located at data/zomato_restaurants.csv
# zomato_df_ind = ...

# YOUR CODE HERE
raise NotImplementedError()
```

```python
# print head it to get an idea of what you've just loaded 
zomato_df_ind.head()
```

<!-- grade_id: cell-c41647c3bbde6448, points: 1, grade: true, solution: false -->
```python
assert(isinstance(zomato_df_ind, pd.DataFrame))
assert(zomato_df_ind.shape == (600, 13))
```

#### 3.3) Get general information about the dataframe

<!-- grade_id: cell-41cd1f7622ecf15e, points: 0, grade: false, solution: true -->
```python
# Print zomato_df_ind info
# reminder: jupyter just prints out the last output, 
# so you don't need to use a print statement

# ...

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-a4e4845a068166a3, points: 0, grade: false, solution: true -->
```python
# How many columns of type "int" exists?

# count_int = ...

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-4d7805dedf30cd4e, points: 1, grade: true, solution: false -->
```python
hash_count = 'e7f6c011776e8db7cd330b54174fd76f7d0216b612387a5ffcfb81e6f0919683'

assert (hash_count == sha256(str(count_int).encode()).hexdigest())   
```

#### 3.4) Preview the top 7 entries

<!-- grade_id: cell-8f9d0a7afe61956f, points: 0, grade: false, solution: true -->
```python
# Print the top 7 entries of zomato_df_ind

# zomato_df_ind_top = ...

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-2b4a36ddf8691172, points: 1, grade: true, solution: false -->
```python
assert (zomato_df_ind_top.shape == (7,13))
```

#### 3.5) Preview the bottom 3 entries

<!-- grade_id: cell-d910c817c876ae9c, points: 0, grade: false, solution: true -->
```python
# Print the bottom 3 entries of zomato_df_ind

# zomato_df_ind_bottom = ...

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-616f1dc1e6e759f5, points: 1, grade: true, solution: false -->
```python
assert (zomato_df_ind_bottom.shape == (3,13))
```

###### 3.6) Use a method to determine all the types os data in zomato_df_ind

<!-- grade_id: cell-99dce98a9e975ca1, points: 0, grade: false, solution: true -->
```python
# Print the type of data in zomato_df_ind

# YOUR CODE HERE
raise NotImplementedError()
```

#### 3.7) Print the number of rows and the number of columns

<!-- grade_id: cell-e4aef4482a81ce78, points: 0, grade: false, solution: true -->
```python
# Get the nr_rows and nr_cols from zomato_df_ind, using a method

#nr_rows = ...
#nr_cols = ...

# YOUR CODE HERE
raise NotImplementedError()

print("There are {} rows and {} columns.".format(nr_rows, nr_cols))
```

<!-- grade_id: cell-5405b4936c60b4b8, points: 2, grade: true, solution: false -->
```python
assert (nr_rows == 600 )
assert (nr_cols == 13 )
```

#### 3.8) Get the array of column names

<!-- grade_id: cell-cd6563ac6cc47e40, points: 0, grade: false, solution: true -->
```python
# Create an array with the column names of zomato_df_ind.
# there is, as always, a method for this 

# zomato_df_columns = ...

# YOUR CODE HERE
raise NotImplementedError()
```

```python
print(zomato_df_columns)
```

The expected output is an array with 'name', 'locality', 'city', 'cuisines', 'establishment','average_cost_for_two', 'all_reviews_count', 'photo_count', 'has_table_booking', 'has_online_delivery', 'aggregate_rating', 'votes','timings'.

<!-- grade_id: cell-9c2f4c475b080292, points: 2, grade: true, solution: false -->
```python
assert (zomato_df_columns[1] == 'locality' )
assert (len(zomato_df_columns) == 13)
```

#### 3.9) Making use of describe() what is the standard deviation value of the votes?

<!-- grade_id: cell-221a101b15ae708a, points: 0, grade: false, solution: true -->
```python
# Print zomato_df_ind describe

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-025c379631d93499, points: 0, grade: false, solution: true -->
```python
# Manually input the standard deviation value of the votes
# yep, that simple. No tricks here :) 

# votes_std = ...

# YOUR CODE HERE
raise NotImplementedError()
```

```python
print("Standard deviation value of the votes is {}.".format(votes_std))
```

<!-- grade_id: cell-c76a75e97007db07, points: 2, grade: true, solution: false -->
```python
assert 594 < votes_std < 595
```