# BLU01 - Exercises Notebook

**Important Note**

When reading files, please use `os.path.join()`. Specially if you are a Windows user!
The grader is running on Linux, reading a file my be running locally for you, but then don't run in the grader because of all the weird backslashes!

```python
import chardet
import hashlib # for grading purposes
import math
import numpy as np
import csv
import pandas as pd
import random
```

## Q1: Use a shell command to keep in 2 variables the first third and the last third of the lines existent in a file





<!-- grade_id: cell-05a3d2245d57305d, points: 0, grade: false, solution: true -->
```python
# Use the file data/exercises/elon_musk.txt
# Start by counting the total lines and add the total to the variable count_total.  
# count_total = ...
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-2e69c6137e73d90c, points: 2, grade: true, solution: false -->
```python
expected_hash = '4ec9599fc203d176a301536c2e091a19bc852759b255bd6818810a42c5fed14a'
assert hashlib.sha256(str(count_total).encode()).hexdigest() == expected_hash
```

<!-- grade_id: cell-14cf2619322a8763, points: 0, grade: false, solution: true -->
```python
# NOT GRADED optional exercise!!
# Now add the first third and the last third of the lines to the variables first_third and last_third. 
# Make it work to any files using count_total/3 for instance in the bash commands
# first_third = ...
# YOUR CODE HERE
raise NotImplementedError()

# last_third = ...
# YOUR CODE HERE
raise NotImplementedError()
```

Test the `first_third` and `last_third` exercise by running the following (ungraded) asserts.

```
assert first_third[-1][0] == 'H'
assert last_third[0][0] == 'N'
```

## Q2: Read a file with specific delimiter

Read file **data/exercises/euribor_interest_rates.csv** into a pandas DataFrame.

First, you should preview the file using a shell command in order to find out the used delimiter, and other properties of this file.

Then, you should use function read_csv to read the data into a DataFrame. The resulting DataFrame should have the last column as index.

<!-- grade_id: cell-942e0590467e20d9, points: 0, grade: false, solution: true -->
```python
# Use a shell command to preview the data
# ! ...
# YOUR CODE HERE
raise NotImplementedError()

# Use function read_csv to read the data into a DataFrame
# df2 = pd.read_csv(...)
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-0c275c9acaa9c74d, points: 2, grade: true, solution: false -->
```python
assert df2.loc[2001, 'Euribor 3 months'] == 3.29
assert set(df2.columns) == {'Euribor 3 months', 'Euribor 6 months', 'Euribor 12 months'}
assert len(df2) == 19
assert df2.index[0] == 1999

expected_hash = 'b68ca0811f132f73edc68e9d3bebb288ef036c1fef8aabe6d2c63a2b6bfa859c'
assert hashlib.sha256(df2.index.name.encode()).hexdigest() == expected_hash
```

## Q3: Read a csv file with problems

Read file **data/exercises/portugal_urban_waste_per_inhabitant.csv** using function `read_csv`. Pay attention to the following:
* you might find some trouble when reading the file, at first, then just ignore the problems ;)
* use the first column as index
* there are some inputs in the file that should be interpreted as NaN, make sure you select the right one when reading the file!

<!-- grade_id: cell-2771b707556c08d2, points: 0, grade: false, solution: true -->
```python
# Read file data/exercises/portugal_urban_waste_per_inhabitant.csv with read_csv
# df3 = pd.read_csv(...)
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-601b354802964776, points: 2, grade: true, solution: false -->
```python
assert np.isnan(df3.loc[1994, 'Urban waste collection per inhabitant (kg/inhabitant)'])
assert df3.loc[2011, 'Selective urban waste collection per inhabitant (kg/inhabitant)'] == 71.4

mean_selective_waste = df3['Selective urban waste collection per inhabitant (kg/inhabitant)'].mean()
assert math.isclose(35.632, mean_selective_waste, rel_tol=1e-3)

expected_hash = 'b68ca0811f132f73edc68e9d3bebb288ef036c1fef8aabe6d2c63a2b6bfa859c'
assert hashlib.sha256(df3.index.name.encode()).hexdigest() == expected_hash
```

## Q4: Repair a csv file after importing
Read the same file **data/exercises/portugal_urban_waste_per_inhabitant.csv** using function `read_csv`. But now, be sure you don't miss any lines with relevant information! 

* use csv module to import everything to a list of lists
* create a df with only 3 meaningful columns, where the `Years` should be index
* replace garbage values with NaN's 
* format the columns with the right type

<!-- grade_id: cell-119041388e17fb72, points: 0, grade: false, solution: true -->
```python
# Read file data/exercises/portugal_urban_waste_per_inhabitant.csv using  
# lines = csv.reader(file)
# YOUR CODE HERE
raise NotImplementedError()
# create a dataframe using the line list with only 3 columns
# df4 = pd.DataFrame(...
# YOUR CODE HERE
raise NotImplementedError()
#replace invalid values with nan (np.nan)
# YOUR CODE HERE
raise NotImplementedError()
#set types per dataframe column
# YOUR CODE HERE
raise NotImplementedError()
#set a new index to the dataframe
# YOUR CODE HERE
raise NotImplementedError()


```

<!-- grade_id: cell-a42299f9047590d5, points: 2, grade: true, solution: false -->
```python
assert np.isnan(df4.loc[1994, 'Urban waste collection per inhabitant (kg/inhabitant)'])
assert df4.loc[2012, 'Selective urban waste collection per inhabitant (kg/inhabitant)'] == 63.3

mean_selective_waste = df4['Selective urban waste collection per inhabitant (kg/inhabitant)'].mean()
assert math.isclose(38.78518518518518, mean_selective_waste, rel_tol=1e-3)

expected_hash = 'b68ca0811f132f73edc68e9d3bebb288ef036c1fef8aabe6d2c63a2b6bfa859c'
assert hashlib.sha256(df4.index.name.encode()).hexdigest() == expected_hash
```

## Q5: Read a JSON file

Read file **data/exercises/portugal_production_of_electricity_gwh.json**.

<!-- grade_id: cell-5c2125479f7e50e5, points: 0, grade: false, solution: true -->
```python
# Read file data/exercises/portugal_production_of_electricity_gwh.json with read_json
# df5 = read_json(...)
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-3f085fa2cdad9319, points: 2, grade: true, solution: false -->
```python
assert len(df5) == 23
assert set(df5.columns) == {'Biomass', 'Geothermal power', 'Hydropower < 10MW', 'Hydropower > 10MW', 
                           'Photovoltaic', 'Total','Total renewable sources','Windpower','Year'}

expected_hash = 'c8226c54f1a24ae847c02a3e7a7d6e9fb44c2f82eb6f1fddcee9092c434e67fb'
assert hashlib.sha256(str(df5.loc[:,'Hydropower > 10MW'].sum()).encode()).hexdigest() == expected_hash
```

## Q6: Read an Excel file

Read file **data/exercises/portugal_gas_emissions_per_year.xlsx** using function read_excel. Pay attention to the following:

* you should grab the table "Series" in sheet "Metadata"
* use column 'Serie' as index
* make sure you keep only the rows with data
* set the variable distinct_scales with the number of ... distinct scales found in the table

<!-- grade_id: cell-017a7c31b0ddc38e, points: 0, grade: false, solution: true -->
```python
# Read file data/exercises/portugal_gas_emissions_per_year.xlsx with read_excel
# df6 = read_excel(...)
# YOUR CODE HERE
raise NotImplementedError()

#distinct_scales = ...
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-91d5f07e40585680, points: 2, grade: true, solution: false -->
```python
assert distinct_scales == 2
assert isinstance(df6, pd.DataFrame)
expected_hash = '60c3ad36f77e7366103fe36a3f551a0cde7d64e26d3102ddb8b953d6208a6006'
assert hashlib.sha256(
        df6.loc[
            df6.index=="Nitrogen oxides", 
            "Measure Unit"][0].encode()
    ).hexdigest() == expected_hash


```

## Q7: Find the encoding of a file

Find the encoding used in file **data/exercises/cities.csv**, using the method that was shown in the Learning Units.

Then, read the data into a DataFrame, using the read_csv method and find the `City` characters that has distance equal to 7503.

<!-- grade_id: cell-bd6944fd63bb38d8, points: 0, grade: false, solution: true -->
```python
# Find the encoding of file data/exercises/mystery_cities.csv
# encoding = ...
# YOUR CODE HERE
raise NotImplementedError()
# Read the file into a DataFrame
# df7 = ...
# YOUR CODE HERE
raise NotImplementedError()

# Find the name of the city with distance = 7503
# distance_found = df7.loc[...]
# YOUR CODE HERE
raise NotImplementedError()

```

<!-- grade_id: cell-687fb83c97e03355, points: 2, grade: true, solution: false -->
```python
assert isinstance(df7, pd.DataFrame)

expected_hash_1 = '969aef39a1d4cb1c5928c774cd7a4e3ccfc064a18fbd43a70193a2631d8a122d'
assert hashlib.sha256(encoding.encode()).hexdigest() == expected_hash_1

expected_hash_2 = '8086d7adc029756c1b9094df64f6e79017dcc5a99c3c5e55b47beede11179a2b'
assert hashlib.sha256(city_found.encode()).hexdigest() == expected_hash_2
```

## Q8: Import a Random Sample of a Big File

Consider the file **data/exercises/world_percentage_of_literacy.tsv**. Let's imagine this file is really huge, with a lot of rows!  Read the file using a random sample of 7 rows. Count the actual lines with `wc` 

<!-- grade_id: cell-fb39107093dd73fb, points: 0, grade: false, solution: true -->
```python
# Read file data/exercises/world_percentage_of_literacy.tsv with wc and save the number of lines
# in lines_in_file. Notice that the header is not a line..
# don't forget to close the opened file
# YOUR CODE HERE
raise NotImplementedError()

# make parameter rows_to_skip equal to the lines you want to skip loading 
# don't forget: 
# - 7 rows should be fecthed)
# - you want to keep the header plus 7 rows in the new dataframe...
# rows_to_skip = ...
# YOUR CODE HERE
raise NotImplementedError()
# Create a df8 file with the sampled values
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-ec15a8be42137a01, points: 2, grade: true, solution: false -->
```python
# getting remaining list of countries, removing rows_to_skip from the file
df_helper = pd.read_csv( 
    'data/exercises/world_percentage_of_literacy.tsv', 
    sep='\t',
    header=0 
)

total_indexes = list(df_helper.index)
for s in rows_to_skip:
    total_indexes.remove(s-1)

assert lines_in_file==194
assert isinstance(df8, pd.DataFrame)
assert list(df8['Country'])==list(df_helper.iloc[total_indexes]['Country'])
assert df8.shape[0]==7

expected_hash = '7902699be42c8a8e46fbbb4501726517e86b22c56a189f7625a6da49081b2451'
assert  hashlib.sha256(str(df8.shape[0]).encode()).hexdigest() == expected_hash

```

## Q9: Loading a Big File

Read file **data/exercises/world_percentage_of_literacy.tsv** using chunks keep only the columns `Country` and `Literacy rate (all)`.
Note that:
* file should be read by chunks of 10 countries
* the missing values should be removed (filtered in each chunk)
* the `Literacy rate (all)` should be converted to type float (in each chunk)
* the index should be incremental starting from 0 (i.e, you don't need to read any column as the index)

In the end calculate the average `Literacy rate (all)`

<!-- grade_id: cell-5106c25a705296e8, points: 0, grade: false, solution: true -->
```python
# Read file data/exercises/world_percentage_of_literacy.tsv
# the chunks should be appended in a list called chunk_arr
# YOUR CODE HERE
raise NotImplementedError()

# df9 should be the final dataframe with concatenated chunks
# Resulting average should go on lit_avg variable
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-f60a22c3465d5b83, points: 2, grade: true, solution: false -->
```python
assert df9.loc[df9['Country']=='World', 'Literacy rate (all)'].values[0] == 86.3
assert df9.dtypes['Country'] == np.object
assert df9.dtypes['Literacy rate (all)'] == np.float

expected_hash = 'f26cd8ca964afd7aaaea0cacb142419676cf9928772f5b5310a036ffdae1586a'
assert hashlib.sha256(str([len(c) for c in chunk_arr]).encode()).hexdigest() == expected_hash

expected_hash = '94f2bba3a658b5642ebbc9b952af45c3db1a185ae0357e4fe7ced784f0c3fe29'
assert hashlib.sha256(str(lit_avg).encode()).hexdigest() == expected_hash
```

## Q10: Calculate average values using chunks and avoiding a complete data frame in memory.

Using chunks, read file **data/exercises/world_percentage_of_literacy.tsv**, avoid incompleted rows and calculate the average of ***Literacy rate (all)*** without loading all data simultaneously. Use a similar approach of the previous question but don't create any dataframe neither any list with chunks;

***Hint: Use the average definition***



<!-- grade_id: cell-57fd7700885246eb, points: 0, grade: false, solution: true -->
```python
# Read file data/exercises/world_percentage_of_literacy.tsv
# the final average should be in the variable final_avg
# You should increment 2 variables in each chunk and use them at the end to calculate final_avg. call them lit_a, lit_b
# YOUR CODE HERE
raise NotImplementedError()

```

<!-- grade_id: cell-0edcf1260d7886c9, points: 2, grade: true, solution: false -->
```python
assert math.isclose(83.24370860927154, final_avg, rel_tol=1e-1)
assert lit_b==151 or lit_a==151
assert int(lit_b) == 12569 or int(lit_a)==12569
```