# SLU09: Model Selection & Overfitting -- Exercises
---

*Exercises are graded unless otherwise indicated.*

```python
import numpy as np
import pandas as pd
from sklearn.linear_model import LinearRegression, LogisticRegression
from utils import generate_test_data
```

## Exercise 1: Detecting bias and variance in the real world (not graded)

For each of the following, identify if they are more likely to be sources of bias or variance:

1. Using very flexible models (e.g., non-parametric, non-linear), such as K-nearest neighbors or decision trees (bias/variance)
2. Using models with simplistic assumptions, such as linear or logistic regressions (bias/variance)
3. Increasing the polynomial degree of our hypothesis function (bias/variance)
4. Ignoring important features (bias/variance)

## Exercise 2: Create training and test datasets (train-test split)

<!-- grade_id: cell-34f35a43586a0e85, points: 0, grade: false, solution: true -->
```python
# import ... 
# YOUR CODE HERE
raise NotImplementedError()

def implement_hold_out_method(X, y, test_size=.4, random_state=0):
    """ 
    Implementing the holdout method, using sklearn.
    
    Args:
        X (pd.DataFrame): a pandas dataframe containing the features
        y (pd.Series): a pandas series containing the target variable
        test_size (float): proportion of the dataset to include in the test set
        random_state (int): the seed used by the random number generator

    Returns:
        X_train (pd.DataFrame): the features for the training examples
        X_test (pd.DataFrame): the features for the test examples
        y_train (pd.Series): target for the training set 
        y_test (pd.Series): target for the test set

    """
    # use train_test_split to create the training and test datasets
    # use named arguments, such as test_size=... and random_state=...
    # X_train, X_test, y_train, y_test = ...
    # YOUR CODE HERE
    raise NotImplementedError()
    
    return X_train, X_test, y_train, y_test
```

<!-- grade_id: cell-d0784c116c6f5dda, points: 3, grade: true, solution: false -->
```python
"""Check that the solution is correct."""

X, y = generate_test_data(m=100, n=4)
X_train, X_test, y_train, y_test = implement_hold_out_method(X, y)

assert X_train.shape == (60, 3)
assert X_test.shape == (40, 3)
assert y_train.shape == (60,)
assert y_test.shape == (40,)

X_train2, X_test2, y_train2, y_test2 = implement_hold_out_method(X, y, test_size=.1)

assert X_train2.shape == (90, 3)
assert X_test2.shape == (10, 3)
assert y_train2.shape == (90,)
assert y_test2.shape == (10,)
```

## Exercise 3: Creating a validation dataset

<!-- grade_id: cell-8a287c4f6e7371cf, points: 0, grade: false, solution: true -->
```python
def implement_validation_dataset(X, y, test_size=.25, val_size=.25, random_state=0):
    """ 
    Implementing the holdout method with validation, using sklearn.
    
    Args:
        X (pd.DataFrame): a pandas dataframe containing the features
        y (pd.Series): a pandas series containing the target variable
        test_size (float): proportion of the dataset to include in the test set
        val_size (float): proportion of the dataset to include in the validation set
        random_state (int): the seed used by the random number generator

    Returns:
        X_train (pd.DataFrame): the features for the training examples
        X_test (pd.DataFrame): the features for the test examples
        X_val (pd.DataFrame): the features of the validation examples
        y_train (pd.Series): target for the training set 
        y_test (pd.Series): target for the test set
        y_val (pd.Series): target for the validation set

    """
    # use train_test_split to split the full dataset into
    # - temp (as in "temporary", as we will split it again) 
    # - test
    # X_temp, X_test, y_temp, y_test = ... (1 line)
    # YOUR CODE HERE
    raise NotImplementedError()
    
    # Great, so now we want to split temp into train and validation. 
    # But what should the "test_size" be when we use the train_test_split function, 
    # so that we end up with 50%-25%-25%?
    # Clue: you can probably use a simple rule for this one
    # the_right_val_size = ... (1 line)
    # YOUR CODE HERE
    raise NotImplementedError()
    
    # use train_test_split to create the train and validation datasets
    # X_train, X_val, y_train, y_val = ... (1 line)
    # YOUR CODE HERE
    raise NotImplementedError()
    
    return X_train, X_test, X_val, y_train, y_test, y_val
```

<!-- grade_id: cell-4cd9ffbea2630f66, points: 3, grade: true, solution: false -->
```python
"""Check that the solution is correct."""
X, y = generate_test_data(m=1000, n=5)
X_train, X_test, X_val, y_train, y_test, y_val = implement_validation_dataset(X, y)

assert X_train.shape == (500, 4)
assert X_test.shape == (250, 4)
assert X_val.shape == (250, 4)
assert y_train.shape == (500,)
assert y_test.shape == (250,)
assert y_val.shape == (250,)


X_train2, X_test2, X_val2, y_train2, y_test2, y_val2 = implement_validation_dataset(X, y, test_size=.1, val_size=.3)

assert X_train2.shape == (600, 4)
assert X_test2.shape == (100, 4)
assert X_val2.shape == (300, 4)
assert y_train2.shape == (600,)
assert y_test2.shape == (100,)
assert y_val2.shape == (300,)
```

## Exercise 4: Implementing K-fold cross-validation

Get the mean cross validation score for linear regression (with 4 folds) and logistic regression (with 6 folds). 

Leave the scoring parameter with the default value. 

<!-- grade_id: cell-b0ed9bb14923e9fd, points: 0, grade: false, solution: true -->
```python
# here is your data 
X_train = pd.read_csv('data/something')
y_train_num = pd.read_csv('data/something_else', header=None)
y_train_cat = y_train_num.apply(lambda x: x < 50).astype('int')

# here are your models
lin_reg = LinearRegression()
log_reg = LogisticRegression(solver='liblinear')

# store the mean cross val scores as lin_reg_result and log_reg_result
# use y_train_num for linear regression
# use y_train_cat for logistic regression
# lin_reg_result = ...
# log_reg_result = ...

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-e89d8d233b70f834, points: 3, grade: true, solution: false -->
```python
np.testing.assert_almost_equal(lin_reg_result, -0.1477, 3)
np.testing.assert_almost_equal(log_reg_result, 0.5333, 3)
```

## Exercise 5: Regularization loss functions

```python
# data 
X = pd.read_csv('data/something')
y = pd.read_csv('data/something_else', header=None)

# model
lr = LinearRegression()

# fit model
lr.fit(X, y)

# get variables for calculating loss functions
y = np.array(y[0])
y_hat = lr.predict(X).reshape(60,)
betas = np.append(lr.intercept_, lr.coef_[0])
```

Compute the $L_1$ and $L_2$ loss functions for the linear regression model trained above.

$$J_{L_1} = \frac{1}{N} \sum_{n=1}^N (y_n - \hat{y}_n)^2 + \lambda_1 \sum_{k=1}^K \left|\beta_k\right|$$

### Exercise 5.1: MSE

First, let's compute the Mean Squared Error part:

<!-- grade_id: cell-add82190e7ce5ca5, points: 0, grade: false, solution: true -->
```python
def mean_squared_error(y, y_hat):
    """
    Args: 
        y : numpy array with shape (num_observations,)
            The targets.
        y_hat : numpy array with shape (num_observations,)
            The predictions.
            
    Returns:
        Mean squared error : float
    """

    # mse = ...
    # YOUR CODE HERE
    raise NotImplementedError()
    
    return mse
```

<!-- grade_id: cell-bc3a09760dfe853f, points: 1, grade: true, solution: false -->
```python
mse = mean_squared_error(y, y_hat)
np.testing.assert_almost_equal(mse, 635.2071, 3)
```

### Exercise 5.2: L1 loss

Now, we can compute the $L_1$ loss

<!-- grade_id: cell-692eb1183bd9f57f, points: 0, grade: false, solution: true -->
```python
def l1_loss(y, y_hat, betas, lamb1):
    """
    Args: 
        y : numpy array with shape (num_observations,)
            The targets.
        y_hat : numpy array with shape (num_observations,)
            The predictions.
        betas : numpy array with shape (num_features+1,)
            The parameters of your regression model. 
            The first value is the intercept and the 
            remaining ones are the feature coefficients.
        lamb1 : float
            The strength of the L1 regularizer.
            
    Returns:
        loss : float
    """
        
    # Compute the mean squared error loss part of the general loss function.
    # Hint: use the function we created above
    # mse = ...
    # YOUR CODE HERE
    raise NotImplementedError()
    
    # Compute the L1 part of the general loss function.
    # l1 = ...
    # YOUR CODE HERE
    raise NotImplementedError()
    
    # Compute the total loss by combining the parts.
    # L = ...
    # YOUR CODE HERE
    raise NotImplementedError()
    
    return L
```

<!-- grade_id: cell-c7a7f5cfc2871021, points: 2, grade: true, solution: false -->
```python
lamb1 = 2
loss = l1_loss(y, y_hat, betas, lamb1)
np.testing.assert_almost_equal(loss, 635.8783, 3)

lamb1_2 = 6
loss2 = l1_loss(y, y_hat, betas, lamb1_2)
np.testing.assert_almost_equal(loss2, 637.2206, 3)
```

### Exercise 5.3: L2 loss

Finally, let's compute the $L_2$ loss:

$$J_{L_2} = \frac{1}{N} \sum_{n=1}^N (y_n - \hat{y}_n)^2 + \lambda_2 \sum_{k=1}^K \beta_k^2$$


<!-- grade_id: cell-c6c5def23f7f4d0a, points: 0, grade: false, solution: true -->
```python
def l2_loss(y, y_hat, betas, lamb2):
    """
    Args: 
        y : numpy array with shape (num_observations,)
            The targets.
        y_hat : numpy array with shape (num_observations,)
            The predictions.
        betas : numpy array with shape (num_features+1,)
            The parameters of your regression model. 
            The first value is the intercept and the 
            remaining ones are the feature coefficients.
        lamb2 : float
            The strength of the L2 regularizer.
            
    Returns:
        loss : float
    """
    
    # Compute the mean squared error loss part of the general loss function.
    # Hint: use the function we created above
    # mse = ...
    # YOUR CODE HERE
    raise NotImplementedError()
    
    # Compute the L2 part of the general loss function.
    # l2 = ...
    # YOUR CODE HERE
    raise NotImplementedError()
    
    # Compute the total loss by combining the parts.
    # L = ...
    # YOUR CODE HERE
    raise NotImplementedError()
    
    return L
```

<!-- grade_id: cell-fecd3123c3630fb1, points: 2, grade: true, solution: false -->
```python
lamb2 = 2
loss = l2_loss(y, y_hat, betas, lamb2)
np.testing.assert_almost_equal(loss, 635.2940, 3)

lamb2_2 = 6
loss2 = l2_loss(y, y_hat, betas, lamb2_2)
np.testing.assert_almost_equal(loss2, 635.4676, 3)
```

## Exercise 6: Regularized linear regression in practice

### Exercise 6.1: Lasso

```python
from sklearn.preprocessing import PolynomialFeatures

# create data
X_temp = pd.read_csv('data/something')
y = pd.read_csv('data/something_else', header=None)
poly = PolynomialFeatures(degree=2)
X = pd.DataFrame(poly.fit_transform(X_temp))
```

<!-- grade_id: cell-041c01ab12c152ea, points: 0, grade: false, solution: true -->
```python
# import ... 
# YOUR CODE HERE
raise NotImplementedError()

def lasso_regression(X, y, lamb):
    """
    Implements lasso (L1) regression, using sklearn
    
    Args: 
        X (pd.DataFrame): a pandas dataframe containing the features
        y (pd.Series): a pandas series containing the target variable
        lamb (float): the strength of the regularizer
            
    Returns:
        lasso : a Lasso regression model fitted to X and y
    """
    
    # instantiate an instance of sklearn's Lasso model
    # use named arguments, such as alpha=... and random_state=...
    # use a random state of 42
    # lasso = ...
    # YOUR CODE HERE
    raise NotImplementedError()
    
    # fit model
    # YOUR CODE HERE
    raise NotImplementedError()
    
    return lasso
```

<!-- grade_id: cell-529b23f08a82851b, points: 2, grade: true, solution: false -->
```python
lamb = 0.1
check_lasso = lasso_regression(X, y, lamb)
np.testing.assert_almost_equal(np.sum(check_lasso.coef_), 0.396970, 6)
np.testing.assert_almost_equal(check_lasso.score(X, y), 0.169879, 6)

lamb2 = 100
check_lasso2 = lasso_regression(X, y, lamb2)
np.testing.assert_almost_equal(np.sum(check_lasso2.coef_), 0.003354, 6)
np.testing.assert_almost_equal(check_lasso2.score(X, y), 0.146606, 6)
```

### Exercise 6.2: Ridge

<!-- grade_id: cell-189ee85987940647, points: 0, grade: false, solution: true -->
```python
# import ... 
# YOUR CODE HERE
raise NotImplementedError()

def ridge_regression(X, y, lamb):
    """
    Implements ridge (L2) regression, using sklearn
    
    Args: 
        X (pd.DataFrame): a pandas dataframe containing the features
        y (pd.Series): a pandas series containing the target variable
        lamb (float): the strength of the regularizer
            
    Returns:
        ridge : a Ridge regression model fitted to X and y
    """
    
    # instantiate an instance of sklearn's Ridge model
    # use named arguments, such as alpha=... and random_state=...
    # use a random state of 42
    # ridge = ...
    # YOUR CODE HERE
    raise NotImplementedError()
    
    # fit model
    # YOUR CODE HERE
    raise NotImplementedError()
    
    return ridge
```

<!-- grade_id: cell-619c183e9b281470, points: 2, grade: true, solution: false -->
```python
lamb = 0.1
check_ridge = ridge_regression(X, y, lamb)
np.testing.assert_almost_equal(np.sum(check_ridge.coef_), 0.401361, 6)
np.testing.assert_almost_equal(check_ridge.score(X, y), 0.169880, 6)

lamb2 = 100
check_ridge2 = ridge_regression(X, y, lamb2)
np.testing.assert_almost_equal(np.sum(check_ridge2.coef_), 0.373489, 6)
np.testing.assert_almost_equal(check_ridge2.score(X, y), 0.169831, 6)
```

### Exercise 6.3: Elastic Net

<!-- grade_id: cell-0dd9fa04c84ff594, points: 0, grade: false, solution: true -->
```python
# import ... 
# YOUR CODE HERE
raise NotImplementedError()

def elasticnet_regression(X, y, lamb, l1_ratio):
    """
    Implements elastic net regression, using sklearn
    
    Args: 
        X (pd.DataFrame): a pandas dataframe containing the features
        y (pd.Series): a pandas series containing the target variable
        lamb (float): total weight of regularization terms
        l1_ratio: the ratio of l1 to l2 loss terms
            
    Returns:
        elasticnet : an Elastic Net regression model fitted to X and y
    """
    
    # instantiate an instance of sklearn's Elastic Net model
    # use named arguments, such as alpha=..., l1_ratio=..., and random_state=...
    # use a random state of 42
    # elasticnet = ...
    # YOUR CODE HERE
    raise NotImplementedError()
    
    # fit model
    # YOUR CODE HERE
    raise NotImplementedError()
    
    return elasticnet
```

<!-- grade_id: cell-018bf67f10314dce, points: 2, grade: true, solution: false -->
```python
lamb = 0.001
l1_ratio = 0.3
check_enet = elasticnet_regression(X, y, lamb, l1_ratio)
np.testing.assert_almost_equal(np.sum(check_enet.coef_), 0.401365, 6)
np.testing.assert_almost_equal(check_enet.score(X, y), 0.169880, 6)

lamb2 = 10
l1_ratio2 = 0.9
check_enet2 = elasticnet_regression(X, y, lamb2, l1_ratio2)
np.testing.assert_almost_equal(np.sum(check_enet2.coef_), 0.039402, 6)
np.testing.assert_almost_equal(check_enet2.score(X, y), 0.159236, 6)
```

```python

```